---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Checking PEP8 standards with flake8 and pylint

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

__PEP8__ stands for *Python Extension Proposal*. The __PEP8__ are used to improve the readability of Python code. There are two major command line tools that help us write __PEP8__ compatible. ```
print("hello, PEP8!")
```

`@instructions`
In the terminal type `pylint name_of_module.py`

`@hint`


`@pre_exercise_code`
```{python}

  
```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Checking PEP8 standards

```yaml
type: VideoExercise
key: 3fe8b70f8f
xp: 50
```

`@projector_key`
d8a5ffe7600ce9865d9ef6fbd2fee35a
